Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 11.2

## Part A.  Currency trading (Slightly simplified version of 15.3-6 from Cormen et al.)

Imagine that you wish to exchange one currency for another. You realize that instead of directly exchanging one currency for another, you might be better off making a series of trades through other currencies, winding up with the currency you want.

Suppose that you can trade n different currencies, numbered $1,2,… ,n$, where you start with currency 1 and wish to wind up with currency $n$. You are given, for each pair of currencies $i$ and $j$ , an exchange rate $r_{ij}$ , meaning that if you start with $d$ units of currency $i$ , you can trade for $dr_{ij}$ units of currency $j$. Note that the total number of trades allowed is limited to $n$.

Assuming there is no commission, propose an algorithm (described in prose) to solve this problem using either a memoization or bottom-up strategy. **Note that we will be working on the Python implementation in class.**

Click [here](https://drive.google.com/open?id=1L8Fjo1Xt8sltab-tz3m91eTuiE9LYWF7) for some example data.

<!-- grade_id: cell-3f8a8808c720baf2, points: 0, grade: true, solution: true -->
This is no different from the rod-cutting problem. We can use a bottom up approach, with the problem being that we try to find a currency numbered t so that when exchanging x to t, then t to y, we get maximum y. This in itself produces two subproblems: find a so that exchanging from x to a, then a to t returns maximum t, and find b so that exchanging from t to b, then b to y returns maximum y, ...

The bottom up approach would solve all the smallest subproblems first. So our algorithm will be something like:

Step 1: Solve the smallest subproblems, which are all possible combinations of exchanging the currency we have with any other currencies (having only one exchange).

Step 2: Solve the next bigger subproblem, which is having two exchanges. All of its subproblems (having 1 exchanges) have already been solved.

Step 3: Repeat until we have all solutions (the final problem is having n exchanges)

Step 4: Print out the optimal solution

## Part B - Money game (Solution)
Consider a row of n coins of values $v_1, v_2,...,v_n$, **where $n$ is even**. We play a game against an opponent by alternating turns. In each turn, a player selects either the first or last coin from the row, removes it from the row permanently, and receives the value of the coin. Determine the maximum possible amount of money we can definitely win if we move first.

For example, consider the game:

\$2, \$10, \$1, \$5

By moving first and playing optimally one can be guaranteed of \$15. The first move is to take \$5. This forces your opponent to take either \$2 or \$1, and then allows you to take \$10. Assume that the opposing player also plays optimally (i.e., minimizing the gain of the first-move player.)

As a hint, we will provide the recurrence for the solution to this problem. Suppose $c(i,j)$ is the maximum possible amount of money the first player can win for the row of coins of value $v_i, v){i+1},...,v_j$ (so to solve our problem, we need to compute $c(1,n)$. Then:

$$c(i,j)=max[v_i+min(c(i+2,j), c(i+1,j-1)), v_j+min(c(i,j-2), c(i+1,j-1))]$$

You are also encouraged to watch [this video](https://www.youtube.com/watch?v=KnP8_L13xW4&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5) if you need some more hints.

## Question 1. 
Complete the following function to solve the game using a bottom-up strategy, assuming that the opposing player also plays optimally. 


<!-- grade_id: cell-4014ae901b66304b, points: 0, grade: false, solution: true -->
```python
def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    n = len(A)
    if n%2 == 1: # if n is an odd number
        return None #return None as per instructions
    elif n==0: #if there are no coins:
        return 0
    else:        
        cache = [[None for i in range(n)] for i in range(n)]
    for diff in range(n): #diff being the length of the sublist, which represents our subproblems
        for i in range(n - diff):
            j = i + diff #j being the end of the sublist
            if diff==0: #if the sublist has only one element, then store itself since that's the only choice
                cache[i][j]=A[i]
            elif diff == 1: #if the sublist has only two element:
                cache[i][j] = max(A[i:j+1]) #store the larger of the two values
            else: #when the sublist has more than two elements:
                cache[i][j] = max(A[i]+min(cache[i+2][j], cache[i+1][j-1]), A[j]+min(cache[i][j-2], cache[i+1][j-1]))
                #this is the given formula
    return cache[0][n-1]
```

<!-- grade_id: cell-f7a5d6c26e95dab2, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

## [Optional]  Question 2. 

Complete the function `print_strategy` to print out a solution to the coin game. Completing the helper function `extended_bottom_up_coin_game` to use in `print_strategy` is optional.

<!-- grade_id: cell-4583755bc245e1e9, points: 0, grade: false, solution: true -->
```python
def print_strategy(A):
    """
    Print coins to pick for the first player
    
    Inputs:
    - A: list of floats, values of the coins, of even length.
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. 
        * max_val is None if the length of list A is odd.
        * max_val is 0 if list A is empty (no coins)
    - out: list of values of the coins that the first player picks in that
    order.
        * out is None if max_val is 0 or None.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ce551c5185e374, points: 0, grade: true, solution: false -->
```python
assert(print_strategy([4,6,5,20,4,6,4,8])[0] == 40)
assert(print_strategy([4,6,5,20,4,6,4,8])[1] == [8, 6, 6, 20])
```